# AdaBoost
For this exercise you will implement AdaBoost from scratch and applied it to a spam dataset. You will be classifying data into spam and not spam. You can call DecisionTreeClassifier from sklearn (with default max_depth=1) to learn your base classifiers.

Here is how you train a decision tree classifier with weights.

`
h = DecisionTreeClassifier(max_depth=max_depth, random_state=0)`

`
h.fit(X, Y, sample_weight=w)
`

```python
import numpy as np
from sklearn.tree import DecisionTreeClassifier
from pathlib import Path
from adaboost import *

from sklearn.tree import DecisionTreeRegressor
from sklearn.metrics import r2_score
from sklearn.model_selection import train_test_split
```

```python
PATH = Path("data")
```

<!-- grade_id: cell-73926460e1f70f54, points: 1, grade: true, solution: false -->
```python
y_test = np.array([1., -1., 1., 1., -1., -1., 1., 1., 1., -1.])
X, Y = parse_spambase_data(PATH/"tiny.spam.train")
for i in range(len(y_test)): assert(y_test[i] == Y[i])
n, m = X.shape
assert(n == 10)
assert(m == 57)
```

<!-- grade_id: cell-4dc8edbf0e11fab8, points: 1, grade: true, solution: false -->
```python
X, Y = parse_spambase_data(PATH/"tiny.spam.train")
trees, weights = adaboost(X, Y, 2)
y_hat_0 = trees[0].predict(X)
assert(len(trees) == 2)
assert(len(weights) == 2)
assert(isinstance(trees[0], DecisionTreeClassifier))
assert(np.array_equal(y_hat_0[:5], [1.,-1.,1, 1, -1]))
```

```python
y_hat_0
```

Output:
```
array([ 1., -1.,  1.,  1., -1., -1., -1.,  1.,  1., -1.])
```

<!-- grade_id: cell-8702186fedbd58e1, points: 2, grade: true, solution: false -->
```python
y_hat_0 = trees[0].predict(X)
assert(np.array_equal(y_hat_0[:5], [1.,-1.,1, 1, -1]))
```

<!-- grade_id: cell-7f28f00061f7b5a5, points: 3, grade: true, solution: false -->
```python
x = np.array([[0, -1], [1, 0], [-1, 0]])
y = np.array([-1, 1, 1])
trees, weights = adaboost(x, y, 1)
pred = adaboost_predict(x, trees, weights)
assert(np.array_equal(pred, y))
```

<!-- grade_id: cell-89126a5a7f8f0e1b, points: 3, grade: true, solution: false -->
```python
X, Y = parse_spambase_data(PATH/"spambase.train")
X_test, Y_test = parse_spambase_data(PATH/"spambase.test")
trees, trees_weights = adaboost(X, Y, 10)
Yhat = adaboost_predict(X, trees, trees_weights)
Yhat_test = adaboost_predict(X_test, trees, trees_weights)
    
acc_test = accuracy(Y_test, Yhat_test)
acc_train = accuracy(Y, Yhat)
print("Train Accuracy %.4f" % acc_train)
print("Test Accuracy %.4f" % acc_test)
assert(np.around(acc_train, decimals=4)==0.9111)
assert(np.around(acc_test, decimals=4)==0.9190)
```

Output:
```
Train Accuracy 0.9111
Test Accuracy 0.9190
```